Gene-Environment Pathways to Cognitive Intelligence and Psychotic-Like Experiences in Children

Author Response

The following is the authors’ response to the previous reviews.

Our comments on the initial eLife assessment

“This study presents a useful inventory of the joint effects of genetic and environmental factors on psychotic-like experiences, and identifies cognitive ability as a potential underlying mediating pathway. The data were analyzed using solid and validated methodology based on a large, multi-center dataset. The claim that these findings are of relevance to psychosis risk and have implications for policy changes are partially supported by the results”

We sincerely appreciate the editor and reviewers for their valuable feedback and their willingness to accommodate our perspectives in the first revision. In this revision, the comments from the reviewers have allowed us to further improve our manuscript. Regarding the eLife assessment, we would like to discuss two points.

Firstly, regarding your point of our “findings are of relevance to psychosis risk…partially supported…”, we want to address that our study is closely related to psychosis risk. Childhood psychotic-like experiences (PLEs) are closely linked to psychotic risk and have been shown to increase the risk of general psychopathology, as mentioned in our Introduction and Discussion.

The reviewers asked for clearer differentiation between PLEs and schizophrenia, which we incorporated in this revision (line 100~111; line 419~430). So, this revised version now clearly points out that findings are relevant primarily to psychosis risk, and only partially relevant to schizophrenia risk.

Secondly, regarding “…implications for policy changes are partially supported…”, we have revised our study’s social contribution more clearly and specifically. Incorporating the comments, we have revised that our study offers an insight to the future studies by showing the importance of integrative approaches, considering multi-factorial neurocognition and psychopathology ranging from genes to environment (line 503~512), rather than offers direct policy implications.

Our collaboration with eLife and the reviewers has proven satisfactory and enriching. The community, coupled with the innovative system and culture established around eLife, has significantly advanced the progression of scientific research. We are privileged to contribute to this endeavor.

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

I am happy with the revisions provided by the authors and I think most of my concerns have been addressed satisfactorily. One remaining concern is the authors' conflation of PLEs and schizophrenia. They stated, for example, that it is necessary to adjust for schizophrenia PGS. Even though studies have found a statistical relationship between schizophrenia PGS and PLEs, this relationship is not very strong (although statistically significant) and other studies have found no relationship. Similarly, having PLEs increases the risk of developing psychosis, but that does not necessarily mean that this risk is substantial or specific. I think this needs more nuance in the manuscript and the term 'schizophrenia' should be used sparsely and very carefully as the paper has focused on PLEs. Otherwise, great work on the revisions, thank you.

Thank you for your comment on the use of PLEs and schizophrenia. We clearly understand the differences between the two and we made relevant corrections throughout the manuscript. In particular, we added that PLEs are not a direct predictor of schizophrenia and corrected any expressions that may imply that PLEs are closely related to schizophrenia in the Introduction.

“Psychotic-like experiences (PLEs), which are prevalent in childhood, indicate the risk of psychosis (van der Steen et al., 2019; Van Os & Reininghaus, 2016). Although they are not a direct precursor of schizophrenia, children reporting PLEs in ages of 9-11 years are at higher risk of psychotic disorders in adulthood (Kelleher & Cannon, 2011; Poulton et al., 2000). PLEs also point towards the potential for other psychopathologies including mood, anxiety, and substance disorders (van der Steen et al., 2019), are linked to deficits in cognitive intelligence (Cannon et al., 2002; Kelleher & Cannon, 2011) and show a stronger association with environmental risk factors during childhood than other internalizing/externalizing symptoms (Karcher, Schiffman, et al., 2021).

Maladaptive cognitive intelligence may act as a mediator for the effects of genetic and environmental risks on the manifestation of psychotic symptoms (Cannon et al., 2000; Keefe et al., 2006; Reichenberg et al., 2005).” (line 100~111)

We also revised any expressions that could be perceived as implying relevance to schizophrenia in the Discussion. “Prior research identifying the mediation of cognitive intelligence focused on either genetic (Karcher, Paul, et al., 2021) or environmental factors (Lewis et al., 2020) alone. Studies with older clinical samples have shown that cognitive deficit may be a precursor for the onset of psychotic disorders (Eastvold et al., 2007; Fett et al., 2020; Vorstman et al., 2015). Our study advances this by demonstrating the integrated effects of genetic and environmental factors on PLEs through the cognitive intelligence in 9-11 years old children. Such comprehensive analysis contributes to assessing the relative importance of various factors influencing children's cognition and mental health, and it can aid future studies designed for identifying health policy implications. Considering the directions and magnitudes of the effects, though the effects of PGS remain significant, aggregated effects of environmental factors account for much greater degrees on PLEs.” (line 419~430)

Reviewer #2 (Recommendations For The Authors):

I thank the authors for addressing most of my comments. I feel the manuscript has already greatly improved.

I have a few more comments.

1. Although I did not make this comment, I find the authors' reply to the following comment by Reviewer #1 unclear: Original comment 'I like that the assessment of CP (cognitive performance) and self-reports PLEs is of good quality. However, I was wondering which 4 items from the parent-reported CBCL (Child Behavior Checklist) were used and how did they correlate with the child-reported PLEs? And how was distress taken into account in the child self-reported PLEs measurement? Which PLEs measures were used?'

The authors' response refers to correlation coefficients, but I think Reviewer #1's inquiry was on more than these correlations.

Thank you for your concern. We think that this comment was referring to our previous manuscript submitted elsewhere. In our initial submission to eLife, we already added the details about the four items from the parent-reported CBCL and how distress was considered in the child self-reported PLEs measurement (Appendix S1, page 48).

1. Regarding the authors' reply that they have 'standardized the use of 'cognitive capacity' - I do not understand what this means. How exactly was this term standardized? In fact, I can find the term 'cognitive capacity' only once and it seemed to have been deleted from the manuscript. This is fine, but it doesn't clearly align with the statement that this term has been standardized.

We apologize for causing such confusion. What we meant was that throughout our revised manuscript, we used the term “cognitive phenotypes” instead of “cognitive capacity”.

1. Regarding my initial comment that 'it needs to be described how cognitive performance was defined in Lee 2018.' - I believe this is still not clarified. The authors write 'CP was measured as the respondent's score on cognitive ability assessments', but it remains unclear what exactly these assessments were.

Thank you for pointing this out. We added that “CP, measured as the respondent's score on cognitive ability assessments of general cognitive function and verbal-numerical reasoning, was assessed in participants from the COGENT consortium and the UK Biobank” (line 204~206).

1. Regarding the authors' reply to my comment 'In the 'Path Modeling' section, please explain what 'factors and components' concretely refer to. How is this different from a standard SEM with latent factors?'

I can see that the authors explained 'components' (=the weighted sum of observed variables), but please also add what you mean by 'factors' - and how these are different from 'components' (line 284). Furthermore, I don't think it is correct that SEMs can only model latent factors, but not components (=measured variables). I also cannot see how using a weighted sum of observed variables controls more effectively for bias in estimation than latent factors. However, even though I do have some knowledge on this method, I'm not an expert and would appreciate the authors, other reviewer and/or editor to weigh in on this point.

Thank you for pointing this out. We added that latent factors are indirectly measured indicators that explain the covariance among observed variables (line 263~271). We also added that standard SEM method using latent factors assumes that observed variables within each construct share a common underlying factor, but if this assumption is not met, then the standard SEM method cannot effectively control for biases. This is the reason why the IGSCA method, which addresses this limitation by allowing for use of both composite and latent factors as constructs.

“Standard SEM using latent factors (i.e., indirectly measured indicators that explain the covariance among observed variables) to represent indicators such as PGS or family SES relies on the assumption that observed variables within each construct share a common underlying factor. If this assumption is violated, standard SEM cannot effectively control for estimation biases. The IGSCA method addresses this limitation by allowing for the use of composite indicators (i.e., components)—defined as a weighted sum of observed variables—as constructs in the model, more effectively controlling bias in estimation compared to the standard SEM. During estimation, the IGSCA determines weights of each observed variable in such a way as to maximize the variances of all endogenous indicators and components.” (line 263~271)

1. I overall disagree with the authors' following statement 'It has been suggested from prior studies that these variables (PGS, family SES, neighborhood SES, positive family and school environment, and PLEs) are less likely to share a common factor', but I appreciate the authors' argument.

Thank you for your comment. To make clarify our statement in the manuscript, we changed the sentence to “Considering that the observed variables of the PGSs, family SES, neighborhood SES, positive family and school environment, and PLEs are evaluated as a composite index by prior research, the IGSCA method can mitigate bias more effectively by representing these constructs as components” (line 274~277).

1. Regarding 'genetic ethnicity': please describe your methods on how this was defined.

Genetic ethnicity was defined as the genetic ancestry of participants, which is included as one of observations in the original ABCD Study data. To avoid further confusion, we corrected ‘genetic ethnicity’ to ‘genetic ancestry’ throughout the manuscript.

1. Regarding 'a more direct genetic predictor of PLEs' - I still don't understand what the contrast is here. More direct than what else?

The description was unclear; we removed it from our manuscript.

1. Regarding the factor loadings in Figure 3: I don't understand how deprivation loads positively on 'low neighborhood SES', but poverty loads negatively. Shouldn't they both show the same direction of effect/loading on neighbourhood SES, while 'years of residency' should show the opposite direction (i.e., deprivation and poverty = risk, while years of residency = protective)? Are these unexpected loadings?

The authors did not yet respond to this point: 'Please also add the autocorrelations between the 3 PLE measures. I assume these were also modelled statistically, given the strong correlations between time points?' Were these correlations not modelled? Why not?

Figure 3B is still unclear. Was intelligence included here? What is the difference between Figure 3A and B? The legend suggests that 3B shows the indirect effects, but figure 3B looks like a direct effect, while 3A seem to show the indirect effect.

The reviewer’s confusion resulted from our incorrect description. The factor loadings of low neighborhood SES were marked incorrectly. The loading for ‘years of residence’ and ‘poverty’ should be switched: -0.3648 for ‘years of residence’ and +0.877 for ‘poverty’. This was a mistake when we were applying factor loadings in the Figure. We thank you for pointing this out.

We apologize for missing your point on autocorrelation. Adding autocorrelations between the three PLEs is unrelated to our research goal. In this paper, we investigated how genetic and environmental factors explain the variations in PLEs between participants, regardless of changes over time. Since we used PLEs of multiple follow-ups to ensure that the results are robust irrespective of the timing of PLE measurements, taking autocorrelation into account is not necessary.

The decision to add autocorrelation, which involves using the outcome variable at time (t-1) as a predictor for the outcome variable at time t, depends on the research focus. If your interest lies in explaining inter-individual variation in the rate of change in PLEs over a one-year period, then autocorrelation should be controlled for (typically, predictors measured at different time points are used in such cases). However, this was not the focus of this paper, which is why we did not apply autocorrelation in the SEM analysis.

We apologize for the confusion between Figure 3A and 3B. To clarify, we added titles in the figure images as “Direct effects” and “Indirect effects”. We also changed the legend as well.

“A. Direct pathways from PGS, high family SES, low neighborhood SES, and positive environment to cognitive intelligence and PLEs. Standardized path coefficients are indicated on each path as direct effect estimates (significance level *p<0.05). B. Indirect pathways to PLEs via intelligence were significant for polygenic scores, high family SES, low neighborhood SES, and positive environment, indicating the significant mediating role of intelligence.” (line 968~973)

Figure 3A shows direct effects: i.e., the coefficients of paths from PGS, family SES, neighborhood SES, and positive environment to intelligence and PLEs, as well as the coefficient of paths from intelligence to PLEs. This is why Figure 3A shows colored arrows starting from PGS, family and neighborhood SES, and positive environment towards intelligence and PLEs, as well as the arrows from intelligence to PLEs. On the other hand, in Figure 3B, the colored arrows staring from PGS, family and neighborhood SES, and positive environment goes through intelligence, and heads towards PLEs. This was meant to show that the indirect effects shown in Figure 3B indicate the specific effects of PGS, family SES, neighborhood SES, and positive environment on PLEs mediated by intelligence.

In short, Figure 3 can be seen as a diagram drawn from Table 2: direct effects of the genetic and environmental variables on intelligence and PLEs, and direct effects of intelligence on PLEs are shown in Figure 3A; indirect effects of genetic and environmental variables on PLEs mediated by intelligence are shown in Figure 3B.

1. Regarding Supporting Information tables: to make these more digestible, I suggest using Excel and adding one table per sheet with a clear title and legend, indicating what each table shows. For example, Table S1 has 9(?) different subsections, all called the same (Linear Mixed Model: Multiethnic). It is not clear how each subsection differs from the others. Separate tables in separate excel sheets might be easier.

Also, I think two decimal points might be good enough, enhancing readability of these tables.

Thank you for your suggestion. We moved the supplementary tables into an external Excel file, with each sheet showing different tables, as well as titles, legends, and clear subsections.

1. Regarding reporting exact p-values in Table 2: I don't understand. At the moment, categorical significance statements are reported. Were these not based on exact p-values (or how else was it decided if a finding was significant at a 0.05 (?) significance level).

Either remove the significance column completely (as p-values cannot be estimated due to non-normality) or specify exactly/clarify what this column shows and this was derived.

We apologize for the confusion. In Table 2, we checked the significance of each path using 95% confidence intervals with 5,000 bootstrapping iterations. Since 95% confidence intervals that does not include zero is equivalent to p-values below 0.05 significance level, we believe this is an appropriate alternative for reporting the significance of each path in the SEM model.

We specified the reason why we were not able to calculate exact p-values (clean copy: line 299~303). “As a trade-off for obtaining robust nonparametric estimates without distributional assumptions for normality, the IGSCA method does not return exact p-values (Hwang, Cho, Jung, et al., 2021). As a reasonable alternative, we obtained 95% confidence intervals based on 5,000 bootstrap samples to test the statistical significance of parameter estimates.”

Changes from the earlier version to the revised version

Abstract

Abstract

Background

In children, psychotic-like experiences (PLEs) are related to risk of psychosis, schizophrenia, and other mental disorders. Maladaptive cognitive functioning is a well-known risk factor and early marker for psychosis, schizophrenia, and other mental disorders. Since cognitive functioning is linked to various genetic and environmental factors during development, we hypothesize that it mediates the effects of those factors on childhood PLEs. Using large longitudinal data, we tested the relationships of genetic and environmental factors (such as familial and neighborhood environment) with cognitive intelligence and their relationships with current and future PLEs in children.

Methods

To estimate associations against potential confounding bias, we leveraged large-scale multimodal data of 6,602 children (aged 9-10 years old; 47.15% females; 5,211 European-ancestry) from the Adolescent Brain and Cognitive Development Study. Linear mixed model and a novel structural equation modeling (SEM) method that allows estimation of both components and factors were used to estimate the joint effects of cognitive phenotypes polygenic scores (PGSs), familial and neighborhood socioeconomic status (SES), and supportive environment on NIH Toolbox cognitive intelligence and PLEs. We adjusted for ethnicity (genetically defined), schizophrenia PGS, and additionally unobserved confounders (using computational confound modeling).

Results

We identified that lower cognitive intelligence and higher PLEs correlated significantly with several genetic and environmental variables: i.e., lower PGSs for cognitive phenotypes, lower familial SES, lower neighborhood SES, lower supportive parenting behavior, and lower positive school environment. In SEM, lower cognitive intelligence significantly mediated the genetic and environmental influences on higher PLEs (Indirect effects of PGS: β range=-0.0355~ -0.0274; Family SES: β range=-0.0429~ -0.0331; Neighborhood SES: β range=0.0126~ 0.0164; Positive Environment: β range=-0.0039~ -0.003). Supportive parenting and a positive school environment had the largest total impact on PLEs (β range=-0.152~ -0.1316) than any other genetic or environmental factors.

Conclusions

Our results reveal the role of genetic and environmental factors on children’s PLEs via its negative impact on cognitive intelligence. Our findings have policy implications in that improving the school and family environment and promoting local economic development might be a way to enhance cognitive and mental health in children.

Key Points

While existing research shows the association between cognitive decline and PLEs, the genetic and environmental pathways to cognitive intelligence and psychotic risk in children remain unclear.

We identified the significant role of genetic and environmental factors (family, neighborhood, and school) on children’s PLEs via a negative impact on cognitive intelligence.

Leveraging large samples with multimodal longitudinal data and advanced computational modeling for adjustment of observed/unobserved confounding bias, our results underscore the importance of incorporating socioeconomic policies into children’s cognitive and mental health programs.

Introduction

Childhood is the critical developmental period in human life. Cognitive intelligence and mental health in this period significantly impact key life outcomes at later ages, including academic performance, economic productivity, physical health, intelligence, and psychopathology (Shonkoff, 2012; Walker et al., 2021). Literature shows the significant impact of social adversities on cognitive ability and mental health in early childhood. Lower family socioeconomic status (SES), particularly household income, is linked to lower neurocognitive ability and higher risk of psychopathology in childhood (Hair et al., 2015; Noble et al., 2015; Peverill et al., 2021; Tomasi & Volkow, 2021; Weissman et al., 2018).

Additional to family SES, the importance of neighborhood social environment on children’s neurocognitive ability has been also emphasized (Gard et al., 2021; Tooley et al., 2020). Adverse neighborhood environment, such as the percent of families below poverty line, low education levels, and exposure to violence, is associated with lower cognitive performance and a greater risk for psychosis in children (Butler et al., 2018; Karcher, Schiffman, et al., 2021; Rakesh et al., 2021; Taylor et al., 2020). Conversely, as protective factors against familial and neighborhood socioeconomic challenges, supportive parenting (Brody et al., 2017; Brody et al., 2019; Holmes et al., 2018; Luby et al., 2012; Luby et al., 2016) and positive school environment (Gard et al., 2021; Piccolo et al., 2019; Rakesh et al., 2021) have been highlighted to improve child cognition and mental health.

Earlier version: Psychotic-like experiences (PLEs), which are prevalent in childhood and often precede the onset of schizophrenia (Howes & Murray, 2014), indicate not only the risk of psychosis (van der Steen et al., 2019; Van Os & Reininghaus, 2016) but also point towards the potential for other psychopathologies including mood, anxiety, and substance disorders (Poulton et al., 2000; van der Steen et al., 2019). PLEs are linked to a reduced cognitive intelligence (Cannon et al., 2002; Kelleher & Cannon, 2011) and show a stronger association with environmental risk factors during childhood than other internalizing/externalizing symptoms (Karcher, Schiffman, et al., 2021).

Revised version: Psychotic-like experiences (PLEs), which are prevalent in childhood, indicate the risk of psychosis (van der Steen et al., 2019; Van Os & Reininghaus, 2016). Although they are not a direct precursor of schizophrenia, children reporting PLEs in ages of 9-11 years are at higher risk of psychotic disorders in adulthood (Kelleher & Cannon, 2011; Poulton et al., 2000). PLEs also point towards the potential for other psychopathologies including mood, anxiety, and substance disorders (van der Steen et al., 2019), are linked to deficits in cognitive intelligence (Cannon et al., 2002; Kelleher & Cannon, 2011) and show a stronger association with environmental risk factors during childhood than other internalizing/externalizing symptoms (Karcher, Schiffman, et al., 2021).

Earlier version: Maladaptive cognitive intelligence may precede the manifestation of psychotic symptoms (Cannon et al., 2000; Keefe et al., 2006; Reichenberg et al., 2005). An integrated theoretical model of schizophrenia posits that anomalous neurodevelopment, influenced by genetic factors, combined with disrupted cognitive processes resulting from socioenvironmental adversity, may eventually give rise to the positive symptoms of schizophrenia, relevant to PLEs (Garety et al., 2001; Howes & Murray, 2014). Family studies show a decline in cognitive intelligence preceding psychotic symptoms is related to genetic risk (Cosway et al., 2000; Curtis et al., 2001). In more recent studies these associations have led to the model positing that cognitive intelligence mediates the genetic risk for psychopathology and PLEs (Karcher, Paul, et al., 2021; Pat et al., 2021). Moreover, there is evidence of a shared genetic component between psychotic disorders and cognitive phenotypes, such as educational attainment (Escott-Price et al., 2020; Le Hellard et al., 2016).

Revised version: Maladaptive cognitive intelligence may act as a mediator for the effects of genetic and environmental risks on the manifestation of psychotic symptoms (Cannon et al., 2000; Keefe et al., 2006; Reichenberg et al., 2005). Abnormal neurodevelopment, influenced by genetic factors, combined with disrupted cognitive processes resulting from socioenvironmental adversity, may eventually give rise to the positive symptoms of schizophrenia, relevant to PLEs (Garety et al., 2001; Howes & Murray, 2014). Family studies show a decline in cognitive intelligence preceding psychotic symptoms is related to genetic risk (Cosway et al., 2000; Curtis et al., 2001). In more recent studies these associations have led to the model positing that cognitive intelligence mediates the genetic risk for psychopathology and PLEs (Karcher, Paul, et al., 2021; Pat et al., 2021).

To minimize potential bias in the estimates of environmental effects, it is crucial to adjust for genetic confounding (Sariaslan et al., 2016), given the substantial genetic influence on intelligence (Bouchard & McGue, 1981; Deary et al., 2006; Plomin & Spinath, 2004) and PLEs (Bentall & Fernyhough, 2008; Maxwell et al., 2023). Recent advances in genetics have led to the development of the polygenic score (PGS) approach: a computational method to estimate the genetic loading for a complex trait using statistical associations of each single-nucleotide polymorphism (SNP) identified by genome-wide association studies (GWAS) (Choi et al., 2020). Particularly, PGS for two related but distinct phenotypes—cognitive performance and educational attainment—holds significant importance. As the two most frequently used proxies of cognitive intelligence in genetic studies, PGSs for cognitive performance and educational attainment are positively correlated with intelligence, educational attainment, income, self-rated health, and height (Judd et al., 2020; Lee et al., 2018; Okbay et al., 2022; Selzam et al., 2019). Furthermore, the PGS for educational attainment is associated with a wide range of biological and social outcomes, including brain morphometry (Judd et al., 2020; Karcher, Paul, et al., 2021), psychopathologies such as PLEs, autism, depression, Alzheimer’s disease, neuroticism (Karcher, Paul, et al., 2021; Okbay et al., 2022), cognitive decline (Joo et al., 2022; Karcher, Paul, et al., 2021; Ritchie et al., 2020), BMI, time spent watching television, geographic residence (Abdellaoui et al., 2022), and wealth inequality (Barth et al., 2020). Similar to the terms used in prior research, we will collectively refer to these two PGSs of focus as ‘cognitive phenotypes PGSs’ throughout this paper (Joo et al., 2022; Okbay et al., 2022; Selzam et al., 2019). An important gap in the literature is the lack of integrated assessment of the effects of genetic and environmental factors at multiple levels (e.g., familial vs neighborhood) to dissect the genetic and environmental effects underlying abnormal cognitive intelligence and the PLEs. Addressing this with large multimodal data will allow for a more complete understanding of the factors related to the development of PLEs.

In this study, we systematically explore the longitudinal trajectories of genetic and environmental influences on PLEs, mediated through cognitive intelligence. Toward this goal, we firstly assess the associations of cognitive phenotype PGSs, family and neighborhood SES, and supportive environment with children’s cognitive intelligence and longitudinally measured PLEs. To maintain robustness of our assessment, we employed statistical and computational approaches to carefully consider potential confounding. We then test the mediating effect of cognitive intelligence on the relationship between genetic and environmental factors and PLEs. Our investigation traces these effects from the baseline and through the 1-year and 2-year follow-ups, providing a nuanced understanding on the role of cognitive phenotype polygenic scores, family socioeconomic status, neighborhood socioeconomic status, and positive family and school environments in shaping PLEs in children aged 9-10 years.

Materials and Methods

Study Participants

Earlier version: We used the multi-modal genetic and environmental data of 11,878 preadolescent children aged 9-10 years old collected from 21 research sites of the Adolescent Brain Cognitive Development (ABCD) Study, one of the largest longitudinal studies for children’s neurodevelopment in the United States. We analyzed the baseline, first year, and second year follow-up datasets included in ABCD Release 4.0, downloaded on January 25, 2022. After k-nearest neighbor imputation of missing values of covariates (categorical variables: sex, genetic ethnicity, marital status of the caregiver, ABCD research sites; continuous variables: age, BMI, family history of psychiatric disorders; 4.67% of total observations imputed) using the R package VIM (Kowarik & Templ, 2016), we removed participants with missing data on study variables (missing genotype: N=3,260; follow-up observations: N=1,180; neighborhood information: N=694; cognitive intelligence tests: N=126; PLEs: N=5; positive environment: N=11). The final samples included 6,602 multiethnic children, which comprised 890 of African ancestry (13.48%), 91 of East Asian ancestry (1.38%), 5,211 of European ancestry (78.93%), 229 of Native American ancestry (3.47%), and 181 not specified (2.74%).

Revised version: We used the multi-modal genetic and environmental data of 11,878 preadolescent children aged 9-10 years old collected from 21 research sites of the Adolescent Brain Cognitive Development (ABCD) Study, one of the largest longitudinal studies for children’s neurodevelopment in the United States. We analyzed the baseline, first year, and second year follow-up datasets included in ABCD Release 4.0, downloaded on January 25, 2022. After k-nearest neighbor imputation of missing values of covariates (categorical variables: sex, genetic ancestry, marital status of the caregiver, ABCD research sites; continuous variables: age, BMI, family history of psychiatric disorders; 4.67% of total observations imputed) using the R package VIM (Kowarik & Templ, 2016), we removed participants with missing data on study variables (missing genotype: N=3,260; follow-up observations: N=1,180; neighborhood information: N=694; cognitive intelligence tests: N=126; PLEs: N=5; positive environment: N=11). The final samples included 6,602 multiethnic children, which comprised 890 of African ancestry (13.48%), 91 of East Asian ancestry (1.38%), 5,211 of European ancestry (78.93%), 229 of Native American ancestry (3.47%), and 181 not specified (2.74%).

Data

NIH Toolbox Cognitive Performance

Children’s neurocognitive abilities were assessed using the NIH Toolbox Cognitive Battery, which has seven cognitive instruments for examining executive function, episodic memory, language abilities, processing speed, working memory, and attention (Thompson et al., 2019). We utilized baseline observations of uncorrected composite scores of fluid intelligence (Dimensional Change Card Sort Task, Flanker Test, Picture Sequence Memory Test, List Sorting Working Memory Test), crystallized intelligence (Picture Vocabulary Task and Oral Reading Recognition Test), and total intelligence (all seven instruments) provided in the ABCD Study dataset.

Psychotic-Like Experiences (PLEs)

Baseline and 1- and 2-year follow-up of PLEs were measured using the children’s responses to the Prodromal Questionnaire-Brief Child Version. In line with previous research (Karcher et al., 2018; Karcher et al., 2020; Karcher, Schiffman, et al., 2021), we computed Total Score and Distress Score, each indicating the number of psychotic symptoms and levels of total distress. Considering self-reports and parent-reports of psychopathology may differ (Achenbach, 2006), we additionally used parent-rated PLEs derived from four items of the Child Behavior Checklist according to previous studies (Karcher et al., 2018; Karcher et al., 2020; Karcher, Schiffman, et al., 2021). Self-reported PLEs and parent-reported PLEs had significant positive correlation (Pearson’s correlation of baseline year: r = 0.095~0.0989, p<0.0001; 1-year follow-up: r = 0.1322~0.1327, p<0.0001; 2-year follow-up: r = 0.1569~0.1632, p<0.0001).

Polygenic Scores (PGS)

Earlier version: To investigate the aggregated effect of genetic components, we estimated PGS of two representative cognitive phenotypes for each participant: educational attainment (EA) and cognitive performance (CP) (Choi et al., 2020). We used the summary statistics released from a GWAS (Lee et al., 2018) of European-descent individuals for EA (n=1,131,881) and CP (n=257,841). EA was measured as the years of schooling; CP was measured as the respondent’s score on cognitive ability assessments. To construct PGS of schizophrenia for sensitivity analyses, we used the summary statistics from the multiple GWAS of European sample (n=65,967; Ruderfer et al. (2018)) and East Asian sample (n=58,140; Lam et al. (2019)). We applied PRS-CSx, a high-dimensional Bayesian regression framework that improves cross-population prediction via continuous shrinkage prior to SNP effect sizes (Ge et al., 2019) (For details, see Appendix S1). The two PGSs for cognitive phenotypes had a positive significant correlation (Pearson’s correlation: r =0.4331, p<0.0001).

Revised version: To investigate the aggregated effect of genetic components, we estimated PGS of two representative cognitive phenotypes for each participant: educational attainment (EA) and cognitive performance (CP) (Choi et al., 2020). We used the summary statistics released from a GWAS (Lee et al., 2018) of European-descent individuals for EA (n=1,131,881) and CP (n=257,841). EA was measured as the years of schooling; CP, measured as the respondent’s score on cognitive ability assessments of general cognitive function and verbal-numerical reasoning, was assessed in participants from the COGENT consortium and the UK Biobank. To construct PGS of schizophrenia for sensitivity analyses, we used the summary statistics from the multiple GWAS of European sample (n=65,967; Ruderfer et al. (2018)) and East Asian sample (n=58,140; Lam et al. (2019)). We applied PRS-CSx, a high-dimensional Bayesian regression framework that improves cross-population prediction via continuous shrinkage prior to SNP effect sizes (Ge et al., 2019) (For details, see Appendix S1). The two PGSs for cognitive phenotypes had a positive significant correlation (Pearson’s correlation: r =0.4331, p<0.0001).

Family-, Neighborhood-, and School-level Environment

We assessed children’s family-level SES with family income, parental education, and family’s financial adversity based on parent self-reporting (Karcher et al., 2020; Taylor et al., 2020; Tomasi & Volkow, 2021). Higher family income and parental education and lower family’s financial adversity denote a higher family SES.

Neighborhood-level SES was assessed using the Area Deprivation Index (ADI), the percentage of individuals below −125% of the poverty level (henceforth “poverty”), and years of residence, which were associated with PLEs in prior research (Karcher, Schiffman, et al., 2021). Higher values of ADI and poverty and fewer years of residence indicate a lower neighborhood SES.

Based on existing literature (Karcher, Schiffman, et al., 2021; Rakesh et al., 2021), we measured the level of positive parenting behavior and positive school environment to assess the effect of positive family and school environment on each individual.

Statistical Modeling

In this study, we employ linear mixed models and a novel structural equation modelling (SEM) method to examine the longitudinal trajectories of genetic and environmental influences on PLEs mediated by cognitive intelligence. We specifically investigate the mediating role of cognitive intelligence within the impacts of cognitive phenotype polygenic scores (PGSs), high family socioeconomic status (SES), low neighborhood SES, and positive family and school environments on PLEs. These influences are examined across three periods of PLEs observations: baseline, 1-year follow-up, and 2-year follow-up (Figure 1).

Study diagram for longitudinal trajectories of genetic and environmental influences on PLEs through cognitive intelligence.

This study examines the mediating role of cognitive intelligence within the effects of cognitive phenotype PGSs, high family SES, low neighborhood SES, and positive family and school environments on PLEs observed at baseline, 1-year follow-up, and 2-year follow-up in children aged 9-10 years. Both direct and indirect effects, as well as total effects, were evaluated for statistical significance.

Linear Mixed Models

Earlier version: Using linear mixed models to minimize potential population stratification from different household environments and geographical locations (Choi et al., 2020), we analyzed the genetic and environmental effects on cognitive intelligence and PLEs. Key variables of interest were PGS, family SES, neighborhood SES, and positive environment. To avoid multicollinearity, CP and EA PGSs were used separately in two other models. As a random factor, the variable indicating ABCD research sites was used. Variables within each model had no signs of multicollinearity (Fox & Monette, 1992) (generalized variance inflation factor < 2.0 for every variable in all models). The child’s sex, age, genetic ethnicity, BMI, marital status of the caregiver, and family history of psychiatric disorders were included as covariates in the model (Karcher, Schiffman, et al., 2021). All continuous variables were standardized (z-scaled) beforehand to get standardized estimates, and the analyses were conducted with the lme4 package in R version 4.1.2. Throughout this paper, threshold for statistical significance was set at p<0.05, with correction for multiple comparisons using the false discovery rate. 95% bootstrapped confidence intervals were obtained with 5,000 iterations.

Revised version: Using linear mixed models to minimize potential population stratification from different household environments and geographical locations (Choi et al., 2020), we analyzed the genetic and environmental effects on cognitive intelligence and PLEs. Key variables of interest were PGS, family SES, neighborhood SES, and positive environment. To avoid multicollinearity, CP and EA PGSs were used separately in two other models. As a random factor, the variable indicating ABCD research sites was used. Variables within each model had no signs of multicollinearity (Fox & Monette, 1992) (generalized variance inflation factor < 2.0 for every variable in all models). The child’s sex, age, genetic ancestry, BMI, marital status of the caregiver, and family history of psychiatric disorders were included as covariates in the model (Karcher, Schiffman, et al., 2021). All continuous variables were standardized (z-scaled) beforehand to get standardized estimates, and the analyses were conducted with the lme4 package in R version 4.1.2. Throughout this paper, threshold for statistical significance was set at p<0.05, with correction for multiple comparisons using the false discovery rate. 95% bootstrapped confidence intervals were obtained with 5,000 iterations.

Path Modeling

Earlier version: To test the plausibility of whether cognitive intelligence may mediate the association between genetic and environmental factors and PLEs, we used an up-to-date SEM method, integrated generalized structured component analysis (IGSCA) (Hwang, Cho, Jung, et al., 2021). This approach is suited to our study using the multi-modal genetic and environmental variables in that it estimates models with both factors and components as statistical proxies for the constructs. Unlike standard SEM with latent factors, the IGSCA method allows for the use of composite indicators (i.e., components)—defined as a weighted sum of observed variables—as constructs in the model, which can more effectively control bias in estimation compared to the standard SEM. During estimation, the IGSCA determines weights of each observed variable in such a way as to maximize the variances of all endogenous indicators and components.

Revised version: To test the plausibility of whether cognitive intelligence may mediate the association between genetic and environmental factors and PLEs, we used an up-to-date SEM method, integrated generalized structured component analysis (IGSCA) (Hwang, Cho, Jung, et al., 2021). This approach is suited to our study using the multi-modal genetic and environmental variables in that it estimates models with both factors and components as statistical proxies for the constructs.

Added in the revised version: Standard SEM using latent factors (i.e., indirectly measured indicators that explain the covariance among observed variables) to represent indicators such as PGS or family SES relies on the assumption that observed variables within each construct share a common underlying factor. If this assumption is violated, standard SEM cannot effectively control for estimation biases. The IGSCA method addresses this limitation by allowing for the use of composite indicators (i.e., components)—defined as a weighted sum of observed variables— as constructs in the model, more effectively controlling bias in estimation compared to the standard SEM. During estimation, the IGSCA determines weights of each observed variable in such a way as to maximize the variances of all endogenous indicators and components.

Earlier version: We considered path-analytic relationships among the six key constructs: cognitive phenotypes PGSs, family SES, neighborhood SES, positive family and school environment, general intelligence, and PLEs. Prior research indicates that the relevant observed variables of the PGSs, family SES, neighborhood SES, positive family and school environment, and PLEs tend not to have a common latent factor and are consequently evaluated as a composite index in analyses. Notably, investigations carried out by Judd et al. (2020) and Martin et al. (2015) utilized composite indicators to examine the genetic influence on educational attainment and ADHD. Moreover, socioenvironmental influences are often treated as composite indicators as highlighted in Judd et al. (2020). When considering the psychosis continuum, studies like that of van Os et al. (2009) postulate that psychotic disorders are likely underpinned by a multiplicity of background factors rather than a single common factor. This perspective is substantiated by a multitude of prior research that deploys composite indices for the measurement of psychotic symptoms. For these reasons, we statistically represented these constructs as the weighted sums of their observed variables or components (Cho et al., 2021). On the other hand, we represented general intelligence as a common factor that determines the underlying covariance pattern of fluid and crystallized intelligence, based on the classical g theory of intelligence (Jensen, 1998; Spearman, 1904).

Revised version: We assessed path-analytic relationships among the six key constructs: cognitive phenotypes PGSs, family SES, neighborhood SES, positive family and school environment, general intelligence, and PLEs. Considering that the observed variables of the PGSs, family SES, neighborhood SES, positive family and school environment, and PLEs are evaluated as a composite index by prior research, the IGSCA method can mitigate bias more effectively by representing these constructs as components. Notably, investigations carried out by Judd et al. (2020) and Martin et al. (2015) utilized composite indicators to examine the genetic influence on educational attainment and ADHD. Moreover, socioenvironmental influences are often treated as composite indicators as highlighted in Judd et al. (2020). When considering the psychosis continuum, studies like that of van Os et al. (2009) postulate that psychotic disorders are likely underpinned by a multiplicity of background factors rather than a single common factor. This perspective is substantiated by a multitude of prior research that deploys composite indices for the measurement of psychotic symptoms. For these reasons, we statistically represented these constructs as the weighted sums of their observed variables or components (Cho et al., 2021). On the other hand, we represented general intelligence as a common factor that determines the underlying covariance pattern of fluid and crystallized intelligence, based on the classical g theory of intelligence (Jensen, 1998; Spearman, 1904).

Earlier version: The IGSCA model included the same covariates used in the linear mixed model as well as the ABCD research site as an additional covariate. We applied GSCA Pro 1.1 (Hwang, Cho, & Choo, 2021) to fit the IGSCA model to the data and checked the model’s goodness-of-fit index (GFI) (Jöreskog & Sörbom, 1986), standardized root mean square residual (SRMR), and total variance of all indicators and components explained (FIT) to assess its overall goodness-of-fit. Ranging from 0 to 1, a larger FIT value indicates more variance of all variables is explained by the specified model (e.g., FIT=0.50 denotes that the model explains 50% of the total variance of all variables) (Hwang, Cho, & Choo, 2021). The rules-of-thumb cutoff criteria in IGSCA is GFI≥0.93 and SRMR<0.08 for an acceptable fit (Cho et al., 2020). Finally, we conducted conditional process analyses to investigate further the indirect and total effects of the constructs in the model. We obtained 95% confidence intervals based on 5,000 bootstrap samples to test the statistical significance of parameter estimates.

Revised version: The IGSCA model included the same covariates used in the linear mixed model as well as the ABCD research site as an additional covariate. We applied GSCA Pro 1.1 (Hwang, Cho, & Choo, 2021) to fit the IGSCA model to the data and checked the model’s goodness-of-fit index (GFI) (Jöreskog & Sörbom, 1986), standardized root mean square residual (SRMR), and total variance of all indicators and components explained (FIT) to assess its overall goodness-of-fit. Ranging from 0 to 1, a larger FIT value indicates more variance of all variables is explained by the specified model (e.g., FIT=0.50 denotes that the model explains 50% of the total variance of all variables) (Hwang, Cho, & Choo, 2021). The rules-of-thumb cutoff criteria in IGSCA is GFI≥0.93 and SRMR≤0.08 for an acceptable fit (Cho et al., 2020). Finally, we conducted conditional process analyses to investigate further the indirect and total effects of the constructs in the model. As a trade-off for obtaining robust nonparametric estimates without distributional assumptions for normality, the IGSCA method does not return exact p-values (Hwang, Cho, Jung, et al., 2021). As a reasonable alternative, we obtained 95% confidence intervals based on 5,000 bootstrap samples to test the statistical significance of parameter estimates.

Sensitivity Analyses

Earlier version: To ensure robustness of the main analyses results, we conducted multiple sensitivity analyses. As the European-descent-based GWAS was used for constructing PGS, we reran the main analyses using participants of European ancestry (n=5,211) to adjust for ethnic confounding. Next, we tested effects of gene x environment interactions on cognitive intelligence and PLEs, respectively. We also assessed whether the effects of cognitive phenotypes PGS in the linear mixed model are significant after adjusting for schizophrenia PGS, given that prior studies reported the significant association of schizophrenia PGS and cognitive deficit in psychosis patients (Shafee et al., 2018) and individuals at-risk of psychosis (He et al., 2021). Lastly, we adjusted for unobserved confounding bias in the linear mixed model, using a recently developed framework for causal inference based on null treatments approach (Miao et al., 2022). Designed to discern causal effects from multiple treatment variables within non-randomized, observational data, the null treatments approach hinges on the assumption that no fewer than half of the confounded treatments exert no causal influence on the outcome. It circumvents the need for prior knowledge regarding which treatments are null and eliminates the necessity for independence among treatments.

Revised version: To ensure robustness of the main analyses results, we conducted multiple sensitivity analyses. As the European-descent-based GWAS was used for constructing PGS, we reran the main analyses using participants of European ancestry (n=5,211) to adjust for ethnic confounding. Next, we tested effects of gene x environment interactions on cognitive intelligence and PLEs, respectively. We also tested the effects of cognitive phenotypes PGS adjusting for schizophrenia PGS, given the association of schizophrenia PGS and cognitive deficit in psychosis patients (Shafee et al., 2018) and individuals at-risk of psychosis (He et al., 2021). Lastly, we adjusted for unobserved confounding bias in the linear mixed model, using a recently developed framework for causal inference based on null treatments approach (Miao et al., 2022). Designed to discern causal effects from multiple treatment variables within non-randomized, observational data, the null treatments approach hinges on the assumption that no fewer than half of the confounded treatments exert no causal influence on the outcome. It circumvents the need for prior knowledge regarding which treatments are null and eliminates the necessity for independence among treatments. Given our model’s inclusion of numerous treatment variables with shared variances due to the presence of unobserved confounders (Abdellaoui et al., 2022; Okbay et al., 2022)—including cognitive phenotypes PGS, family and neighborhood SES, positive family and school environments— we opted to employ this method.

Removed in the revised version: Given our model’s inclusion of numerous treatment variables with shared variances due to the presence of unobserved confounders (Abdellaoui et al., 2022; Okbay et al., 2022)—including cognitive phenotypes PGS, family and neighborhood SES, positive family and school environments— we opted to employ this method.

Results

Demographics

Table 1 presents the demographic characteristics of the final samples. For multiethnic subjects (main analyses, n=6,602), 47.15% were female, and the parents of 70.21% were married. In European ancestry samples (sensitivity analyses, n=5,211), 46.71% were female, and the parents of 77.47% were married. Children of European ancestry showed significantly different marital status (p<0.0001), lower BMI (p<0.0001) and family history of psychiatric disorders (p<0.0001) compared to children of other genetic ancestries. Our linear mixed model and IGSCA analyses were adjusted using sex, age, marital status, BMI, family history of psychiatric disorders, and ABCD research sites as covariates.

Genetic influence on cognitive phenotypes correlates positively with intelligence and negatively with PLEs

As shown in Figure 2, higher PGSs of cognitive capacity phenotypes correlated significantly with higher intelligence (CP PGS: β=0.1113~0.1793; EA PGS: β=0.0699~0.1567). While CP PGS was associated only with lower baseline year Distress Score PLEs (β=-0.0323), EA PGS was associated with lower baseline year and follow-up PLEs of all measures (baseline: β=-0.0518~-0.0519; 1-year: β=-0.0423~-0.043; 2-year: β=-0.036~ −0.0463). No significant correlations were found between CP PGS and follow-up PLEs (p>0.05). The effects of EA PGS were larger on baseline year PLEs than follow-up PLEs. The effect sizes of EA PGS on children’s PLEs were larger than those of CP PGS (EA PGS: β=-0.036~-0.0519; CP PGS: β=-0.0323). (Table S1).

Linear models testing genetic, socioeconomic, and environmental factors associated with cognitive intelligence and PLEs.

Standardized coefficients of a linear mixed model with CP PGS (A, B) and EA PGS (C, D). The analyses included 6,602 samples of multiethnicity. Cognitive intelligence and PLEs correlated with the PGSs, residential disadvantage, positive environment, and family SES in the opposite directions. Error bars indicate 95% bootstrapped confidence intervals with 5,000 iterations. CP and EA denote cognitive performance and education attainment, respectively; PGS, polygenic scores; SES, socioeconomic status; PLEs, psychotic-like experiences; ADI, Area Deprivation Index; Poverty, percentage of individuals below −125% of the poverty level; Years, years of residence. (*p-FDR<0.05)

Family and neighborhood SES correlates positively with intelligence and negatively with PLEs

Parental education associated positively with all types of intelligence (β=0.0699~0.1745) and negatively with baseline year Total and Distress Score PLEs (β=-0.0528~-0.043), 1-year follow-up PLEs (β=-0.0538~-0.0449), and 2-year follow-up PLEs (β=-0.0459~-0.0389). Family income correlated positively with intelligence (β=0.0723~0.1365) and negatively with 2-year follow-up parent-rated PLEs (β=-0.0503~ - 0.0502). Family’s financial disadvantage correlated negatively with baseline year parent-rated PLEs (β=0.0726~0.0728), 1-year follow-up PLEs of all types (β=0.0307~0.0577), and 2-year follow-up PLEs of all types (β=0.0461~0.0581).

The ADI correlated significantly negatively with all types of intelligence (β=-0.0684~-0.054). Additionally, a higher ADI correlated significantly with higher baseline year PLEs (β=0.0587~0.0914), 1-year follow-up PLEs (β=0.0523~0.0613), and 2-year follow-up PLEs (β=0.0397~0.0449).

We found no significant associations of poverty with any of the target variables (p> 0.05). Years of residence correlated significantly with crystallized intelligence (β=0.035~0.0372) and baseline year Total Score PLEs (β=-0.029~-0.0273). (Table S1).

Positive family and school environments correlate positively with intelligence and negatively with the influence of PLEs

Positive parenting behaviors showed significant negative correlations with baseline year PLEs (β=-0.0702~-0.0419), 1-year follow-up PLEs (β=-0.0588~-0.0397), and 2-year follow-up PLEs (β=-0.0623~-0.0356) (Figure 2). Positive school environment was associated positively with total intelligence (β=0.0353~0.0397) and fluid intelligence (β=0.0514~0.0545) and negatively with all three measures of baseline year PLEs (β=-0.1193~-0.0468), 1-year follow-up PLEs (β=-0.1078~-0.04), and 2-year follow-up PLEs (β=-0.1068~-0.0586) (Table S1).

Structural Equation Modeling-Integrated Generalized Structured Component Analysis

Earlier version: The IGSCA model showed that intelligence mediated the effects of genes and environments on the risk for psychosis (PLEs) (Figure 3 and Table 2). Estimated factor loadings of latent factor variable and weights of component variables are presented in Table S2. The model showed a good model fit with a GFI of 0.9735, SRMR of 0.0359, and FIT value of 0.4912 (Cho et al., 2020). Intelligence was under significant direct influences of the cognitive phenotypes PGS (β=0.2427), family SES (β=0.2932), neighborhood SES (β=- 0.1121), and positive environment (β=0.0268). Family SES and positive environment had significant negative direct effects on PLEs of all years. Cognitive phenotypes PGS and neighborhood SES showed no significant direct effects on any of the PLEs (p>0.05).

Revised version: The IGSCA model showed that intelligence mediated the effects of genes and environments on the risk for psychosis (PLEs) (Figure 3 and Table 2). Estimated factor loadings of latent factor variable and weights of component variables are presented in Table S2. The model showed a good model fit with a GFI of 0.9735, SRMR of 0.0359, and FIT value of 0.4912 (Cho et al., 2020). Intelligence was under significant direct influences of the cognitive phenotypes PGS (β=0.2427), family SES (β=0.2932), neighborhood SES (β=-0.1121), and positive environment (β=0.0268). Family SES and positive environment had significant negative direct effects on PLEs of all years. Cognitive phenotypes PGS and neighborhood SES showed no significant direct effects on any of the PLEs (p>0.05). Intelligence significantly mediated the effects of the PGS, family and neighborhood SES, and positive environment on PLEs of all years: the PGS (baseline year: β=-0.035; 1-year: β=-0.0355; 2-year: β=-0.0274), family SES (baseline year: β=-0.0423; 1-year: β=--0.0429; 2-year: β=-0.0331), neighborhood SES (baseline year: β=0.0162; 1-year: β=0.0164; 2-year: β=0.0126), and positive environment (baseline year: β=-0.0039; 1-year: β=-0.0039; 2-year: β=-0.003).

Added in the revised version: Direct/Indirect effects of gene-environment factors to cognitive and PLEs outcomes

Earlier version: Pathway diagram of the structural equation model of gene-environment-intelligence-PLEs pathway: Cognitive phenotypes PGSs, high family SES, low neighborhood SES, positive environment, intelligence, and PLEs. A. Standardized path coefficients are indicated on each path as direct effect estimates (significance level *p<0.05). B. Indirect pathways to PLEs via intelligence were significant for polygenic scores, high family SES, low neighborhood SES, and positive environment, indicating the significant mediating role of intelligence. C. Relative effect sizes of direct and indirect pathways within the total effects on PLEs. The standardized effect sizes of direct pathways are colored within each bar. In A-B, child sex, genetic ethnicity, BMI, marital status, family history of psychiatric disorders, and ABCD research sites were included as covariates. CP PGS and EA PGS denote polygenic scores of cognitive performance and education attainment, respectively; SES, socioeconomic status; PLEs, psychotic-like experiences; Crystallized and Fluid, crystallized and fluid intelligence; ADI, Area Deprivation Index; Poverty, percentage of individuals below -125% of the poverty level; Years, years of residence. Note: * indicates a statistically significant parameter estimate at α=0.05.

Revised version: A. Direct pathways from PGS, high family SES, low neighborhood SES, and positive environment to cognitive intelligence and PLEs. Standardized path coefficients are indicated on each path as direct effect estimates (significance level *p<0.05). B. Indirect pathways to PLEs via intelligence were significant for polygenic scores, high family SES, low neighborhood SES, and positive environment, indicating the significant mediating role of intelligence. C. Relative effect sizes of direct and indirect pathways within the total effects on PLEs. The standardized effect sizes of direct pathways are colored within each bar. In A-B, child sex, genetic ancestry, BMI, marital status, family history of psychiatric disorders, and ABCD research sites were included as covariates. CP PGS and EA PGS denote polygenic scores of cognitive performance and education attainment, respectively; SES, socioeconomic status; PLEs, psychotic-like experiences; Crystallized and Fluid, crystallized and fluid intelligence; ADI, Area Deprivation Index; Poverty, percentage of individuals below −125% of the poverty level; Years, years of residence.

Added in the revised version: Note: * indicates a statistically significant parameter estimate at α=0.05.

Removed in the revised version: Intelligence significantly mediated the effects of the PGS, family and neighborhood SES, and positive environment on PLEs of all years: the PGS (baseline year: β=-0.035; 1-year: β=- 0.0355; 2-year: β=-0.0274), family SES (baseline year: β=-0.0423; 1-year: β=--0.0429; 2- year: β=-0.0331), neighborhood SES (baseline year: β=0.0162; 1-year: β=0.0164; 2-year: β=0.0126), and positive environment (baseline year: β=-0.0039; 1-year: β=-0.0039; 2-year: β=-0.003).

For all observed years, positive environment had largest total effects on PLEs (baseline year: β=-0.152; 1-year: β=-0.1316; 2-year: β=-0.1364), followed by family SES (baseline year: β=-0.1216; 1-year: β=-0.1119; 2-year: β=-0.1164), neighborhood SES (baseline year: β=0.0504; 1-year: β=0.0329; 2-year: β=0.0192), and PGS (baseline year: β=-0.0498; 1-year: β=-0.036; 2-year: β=-0.0365). The total effects of each indicator on PLEs were significant except for those of neighborhood SES (p>0.05).

Sensitivity Analyses

Added in the revised version: As sensitivity analyses, we reran our main analyses with adjustment for ethnic confounding, schizophrenia PGS, and unobserved confounding, respectively. Results of linear mixed model and IGSCA analyses were consistent (Table S3-S7). See Appendix S2 for detailed results.

Removed in the revised version: As sensitivity analyses, we reran our main analyses with adjustment for ethnic confounding, schizophrenia PGS, and unobserved confounding, respectively. Results of linear mixed model and IGSCA analyses were consistent with European ancestry participants (Table S3, S4). Interactions between PGS and positive parenting behavior had significant positive effects on 2-year follow-up Total and Distress Score PLEs (β=0.0343∼0.053) (Table S5). Unlike CP PGS and EA PGS, schizophrenia PGS did not have significant associations with any of the PLEs (p>0.05). Associations of cognitive capacitive PGS, family and neighborhood SES, and positive environment were similar after adjusting for schizophrenia PGS (See Table S6 for details). Lastly, when adjusting for unobserved confounders, CP and EA PGS, positive parenting behavior, and positive school environment remained significantly correlated with intelligence and PLEs of all years (Table S7).

Discussion

This study investigated the relationships of the genetic and environmental influences on the development of intelligence and the PLEs in youth, leveraging genetic data from the large epidemiological samples and a multi-level environmental (socioeconomic) data. Our results support the model that genetic factors (PGS for cognitive phenotypes), socioeconomic conditions, and family and school environments may influence cognitive intelligence in children, and this impact may lead to the individual variability of the current and future PLEs in children. Our analysis with integrated data shows the contributions of genetic and environmental factors, respectively, to cognitive and mental wellness in children, and further provides policy implications to improve them.

Earlier version: Our SEM shows that cognitive intelligence mediates the environmental and genetic influence on the current and future PLEs. The environmental factors (family SES, neighborhood SES, and positive parenting and schooling) and PGS of cognitive phenotypes exhibit significant indirect effects via cognitive intelligence on PLEs. Prior studies identifying the mediation of cognitive intelligence focused on either genetic or environmental factors alone (Karcher, Paul, et al., 2021; Lewis et al., 2020). Our study advances this by demonstrating the specific effects of genetic and environmental factors while controlling for each other. Considering the directions and magnitudes of the effects, though the effects of PGS remain significant, aggregated effects of environmental factors account for much greater degrees on PLEs (potential policy implications is discussed below). The indirect pathway of PGS & environment-intelligence-PLEs aligns with the studies with older clinical samples showing cognitive deficits as a precursor for the onset of psychosis and schizophrenia and their genetic and environmental factors (Eastvold et al., 2007; Fett et al., 2020; Lewis et al., 2020; Vorstman et al., 2015).

Revised version: Our SEM shows that cognitive intelligence mediates the environmental and genetic influence on the current and future PLEs. The environmental factors (family SES, neighborhood SES, and positive parenting and schooling) and PGS of cognitive phenotypes exhibit significant indirect effects via cognitive intelligence on PLEs. Prior research identifying the mediation of cognitive intelligence focused on either genetic (Karcher, Paul, et al., 2021) or environmental factors (Lewis et al., 2020) alone. Studies with older clinical samples have shown that cognitive deficit may be a precursor for the onset of psychotic disorders (Eastvold et al., 2007; Fett et al., 2020; Vorstman et al., 2015). Our study advances this by demonstrating the integrated effects of genetic and environmental factors on PLEs through the cognitive intelligence in 9-11 years old children. Such comprehensive analysis contributes to assessing the relative importance of various factors influencing children’s cognition and mental health, and it can aid future studies designed for identifying health policy implications. Considering the directions and magnitudes of the effects, though the effects of PGS remain significant, aggregated effects of environmental factors account for much greater degrees on PLEs.

Our results of cognitive intelligence mediating the genetic and environmental effects on PLEs may be related to several potential mechanisms. Children raised in higher family SES may have sufficient nutrition and cognitive stimulants, whereas children living in deprived neighborhoods may be exposed to higher rates of crime, air pollution, and substance abuse (Lewis et al., 2020; Marshall et al., 2020; Taylor et al., 2020; Tomasi & Volkow, 2021). Environmental enrichment may be associated with longer periods of neural plasticity (e.g. myelination, maturation of brain circuitry), leading to higher cognitive ability and lower risk of mental disorders like PLEs (Tooley et al., 2021). This may be further linked to the cognitive reserve theory. The theory suggests that genetic influence for cognitive phenotypes and environmental enrichment promotes more efficient, flexible brain networks, which may lead to greater resilience against psychopathology (Stern, 2009). Indeed, prior clinical studies show the linkage between cognitive reserve and psychosis (Amoretti et al., 2018; Leeson et al., 2011).

Earlier version: Our results indicate that genetic influences on cognitive phenotypes are significantly linked to PLEs. While schizophrenia PGS showed no significant associations with PLEs, PGSs for CP and EA were strongly correlated with PLEs (baseline year, 1-year follow-up, and 2-year follow-up). These associations were robust after adjustment for ethnic confounding and unobserved confounders. Cognitive phenotypes PGS generally show higher predictive performance than PGS of any other traits (Lee et al., 2018; Okbay et al., 2022; Plomin & von Stumm, 2018). Genetic variants associated with CP and EA are related to complex traits across the life span, including neuroticism, depressive symptoms, smoking in adulthood, cognitive decline at a later age (Joo et al., 2022), risk for Alzheimer’s disease (Lee et al., 2018; Okbay et al., 2022), brain volume, area, and thickness, as well as psychotic disorders (Karcher, Paul, et al., 2021). Prior gene expression studies suggest that polygenic signals for schizophrenia, bipolar disorder, and EA are significantly enriched in the central nervous system, particularly the cerebellum (Finucane et al., 2018). Our findings emphasize the importance of cognitive phenotypes PGS as a biomarker which not only implicates cognitive traits but also exhibits genetic overlap with the PLEs.

Revised version: Our results indicate that genetic influences on cognitive phenotypes are significantly linked to PLEs. PGSs for CP and EA were strongly correlated with PLEs (baseline year, 1-year follow-up, and 2-year follow-up). These associations were robust after adjustment for schizophrenia PGS, ethnic confounding and unobserved confounders. Cognitive phenotypes PGS generally show higher predictive performance than PGS of any other traits (Lee et al., 2018; Okbay et al., 2022; Plomin & von Stumm, 2018). Genetic variants associated with CP and EA are related to complex traits across the life span, including neuroticism, depressive symptoms, smoking in adulthood, cognitive decline at a later age (Joo et al., 2022), risk for Alzheimer’s disease (Lee et al., 2018; Okbay et al., 2022), brain volume, area, and thickness, as well as psychotic disorders (Karcher, Paul, et al., 2021). Prior gene expression studies suggest that polygenic signals for schizophrenia, bipolar disorder, and EA are significantly enriched in the central nervous system, particularly the cerebellum (Finucane et al., 2018). Our findings emphasize the importance of cognitive phenotypes PGS as a biomarker which not only implicates cognitive traits but also exhibits genetic overlap with the PLEs.

The differing magnitudes of the PGS impact between EA and CP warrant attention. The effects of the EA PGS on the PLEs of all years were 160.68%~371.67% larger than those of CP PGS. This discrepancy may result from that the larger sample size of EA GWAS than that of CP GWAS. Alternatively, the discrepancies in effect sizes may suggest different genetic compositions between EA and CP. Recent literature documents that more than half of the polygenic signal for EA is related to noncognitive and social skills required for successful educational attainment (Demange et al., 2021), whereas CP may rather be linked to cognitive skills. This observation also supports the well-established relationships of the EA PGS with socioeconomic and life-course outcomes (e.g., social mobility (Belsky et al., 2018), voter turnout (Aarøe et al., 2020), BMI, income, time spent watching television, geographic residence (Abdellaoui et al., 2022), and wealth inequality (Barth et al., 2020), which may be influenced by unobserved environmental factors (Young et al., 2019). In our analysis, the utilization of two PGSs a more comprehensive evaluation, contributing to an estimation of the genetic and environmental factors that attempted to minimize confounding bias.

Earlier version: Furthermore, the significant effects of cognitive phenotypes PGS on cognitive intelligence (β=0.0699∼0.1793) remained robust and similar in magnitude after adjusting for genetic ethnicity (β=0.0754∼0.1866) and other (unobserved) confounding (β=0.0546∼0.1776). As we controlled for family-, neighborhood-, and school-level environmental factors and unobserved confounders, our results may be interpreted as significant genetic influences on individual’s cognitive intelligence. This interpretation is supported by a recent study (Isungset et al., 2022): Despite of the socioeconomic differences in Norway (a typical social democratic welfare state) and the US (a typical liberal welfare state), the magnitudes of genetic influence on cognitive intelligence were similar (Norway: β=0.18; US: β=0.17). Cognitive phenotypes PGS is an important genetic factor across the nations and societies. Therefore, analyses omitting the genetic influence may be subject to overestimation of the socioeconomic impact (Plomin & von Stumm, 2018; Sariaslan et al., 2016).

Revised version: Furthermore, the significant effects of cognitive phenotypes PGS on cognitive intelligence (β=0.0699~0.1793) remained robust and similar in magnitude after adjusting for genetic ancestry (β=0.0754~0.1866) and other (unobserved) confounding (β=0.0546~0.1776). As we controlled for family-, neighborhood-, and school-level environmental factors and unobserved confounders, our results may be interpreted as significant genetic influences on individual’s cognitive intelligence. This interpretation is supported by a recent study (Isungset et al., 2022): Despite of the socioeconomic differences in Norway (a typical social democratic welfare state) and the US (a typical liberal welfare state), the magnitudes of genetic influence on cognitive intelligence were similar (Norway: β=0.18; US: β=0.17). Cognitive phenotypes PGS is an important genetic factor across the nations and societies. Therefore, analyses omitting the genetic influence may be subject to overestimation of the socioeconomic impact (Plomin & von Stumm, 2018; Sariaslan et al., 2016).

Earlier version: This study shows that a high SES and positive environment, particularly positive parenting behavior and school environment, is associated with higher intelligence and a lower risk for PLEs in children. While prior research has emphasized the dominant role of family SES (e.g. family income) (Tomasi & Volkow, 2021), our SEM analyses(IGSCA) showed that positive environmental factors such as supportive parenting and schooling have a greater impact on children’s PLEs. Specifically, the effect sizes were the highest in supportive family and school environment, followed by family and neighborhood SES. Even after adjusting for genetic ethnicity and unobserved confounders, the strong associations of positive parenting and schooling with higher intelligence and fewer PLEs remained significant. These findings suggest that interventions that target positive family and school environments may be particularly effective. Recent research supports this notion, showing that interventions that promote supportive parenting and inclusive school environments can improve neurocognitive ability, academic performance, and decrease risk behaviors such as drinking and emotional eating (Brody et al., 2017; Brody et al., 2019; Holmes et al., 2018).

Revised version: This study shows that a high SES and positive environment, particularly positive parenting behavior and school environment, is associated with higher intelligence and a lower risk for PLEs in children. While prior research has emphasized the dominant role of family SES (e.g. family income) (Tomasi & Volkow, 2021), our SEM analyses (IGSCA) showed that positive environmental factors such as supportive parenting and schooling have a greater impact on children’s PLEs. Specifically, the effect sizes were the highest in supportive family and school environment, followed by family and neighborhood SES. Even after adjusting for genetic ancestry and unobserved confounders, the strong associations of positive parenting and schooling with higher intelligence and fewer PLEs remained significant. These findings suggest that interventions that target positive family and school environments may be particularly effective. Recent research supports this notion, showing that interventions that promote supportive parenting and inclusive school environments can improve neurocognitive ability, academic performance, and decrease risk behaviors such as drinking and emotional eating (Brody et al., 2017; Brody et al., 2019; Holmes et al., 2018).

Moreover, our results showed that positive parenting and schooling in baseline year were associated not only with baseline year PLEs but also with PLEs 1-2 years later. This is in line with prior research showing that intervention focused on parenting behavior and school environment have long-lasting positive effects that extend into adulthood and even across generations (Cunha & Heckman, 2007; Hill et al., 2020).

Earlier version: Our findings highlight the potential for preventative strategies that target residential environment, family SES, parenting, and schooling—a comprehensive approach that considers the entire ecosystem of children’s lives—to enhance children’s cognitive ability and mental health. Despite the undeniable genetic influence on PLEs, when we combine the total effect sizes of neighborhood and family SES, as well as positive school environment and parenting behavior (∑ |P| = 0.2718∼0.3242), they considerably surpass the total effect sizes of cognitive phenotypes PGSs (|P| = 0.0359∼0.0502). A holistic and quantitative approach that takes into account the comprehensive ecosystem of family, school, and residential environments may ensure treatment effects and efficient use of resources (Cree et al., 2018; Garner & Yogman, 2021; Shonkoff, 2012). For example, the Health Impact in 5 Years Initiative of the US Centers for Disease Control and Prevention (CDC, 2018) includes 14 evidence-based interventions, such as providing school-based prevention programs, public transportation, home improvement loans, and earned income tax credits, to tackle the social determinants of public health. Our study strengthens the idea that an interdisciplinary science-driven, coordinated approach to intervening in the select environmental factors may allow practical improvements in child development, particularly in those who are at a disadvantage.

Revised version: While policy implications in observational studies like ours might be limited, our findings show the importance of comprehensive approaches considering the entire ecosystem of children’s lives—including residential, family, and school environment—for future research aimed to enhance children’s intelligence and mental health. When we combine the total effect sizes of neighborhood and family SES, as well as positive school environment and parenting behavior (∑ |β| = 0.2718~0.3242), they considerably surpass the total effect sizes of cognitive phenotypes PGSs (|β| = 0.0359~0.0502). It has been suggested that a holistic and quantitative approach that takes into account the comprehensive ecosystem of family, school, and residential environments may ensure policy effects and efficient use of resources (Cree et al., 2018; Garner & Yogman, 2021; Shonkoff, 2012). For example, the Health Impact in 5 Years Initiative of the US Centers for Disease Control and Prevention (CDC, 2018) includes 14 evidence-based interventions, such as providing school-based prevention programs, public transportation, home improvement loans, and earned income tax credits, to tackle the social determinants of public health. Our study strengthens the idea that an interdisciplinary science-driven, coordinated approach to intervening in the select environmental factors may allow practical improvements in child development, particularly in those who are at a disadvantage.

Earlier version: Our study has some limitations. First, due to data availability constraints in the ABCD study, we only utilized baseline observations for NIH Toolbox cognitive intelligence, and we could not test whether PLEs might be a mediator of intelligence. Second, the generalizability of our findings may be limited since most of the participants included in our analysis are from European ancestry. Although the ABCD Study aimed to achieve its representativeness by recruiting from an array of school systems located around each of the 21 research sites, chosen for their diversity in geography, demographics, and socioeconomic status, it is not fully representative of the US population (Compton et al., 2019). Third, the duration of the follow-up period utilized in this study is relatively short (1-year and 2-year follow-up), which may limit the interpretability of our findings for understanding cognitive and psychiatric development during later childhood. Future research could potentially benefit from employing longer follow-up periods, as more follow-up observations are being collected in the ABCD Study. Fourth, while we used a wide range of statistical methods to adjust for confounding bias from observed and unobserved variables (e.g., genetic ethnicity), we did not account for other types of potential bias such as sample selection bias. Fifth, despite a number of causal inference methods used in this study, the ABCD Study is a non-randomized dataset. Given the observational nature of the ABCD Study, interpreting our results as actual causality requires more caution. Finally, we did not include all important environmental variables, such as air pollution (Marshall et al., 2020) and social capital (Krabbendam, 2005), which are not collected in the ABCD study.

Revised version: Our study has some limitations. First, due to data availability constraints in the ABCD study, we only utilized baseline observations for NIH Toolbox cognitive intelligence, and we could not test whether PLEs might be a mediator of intelligence. Second, the generalizability of our findings may be limited since most of the participants included in our analysis are from European ancestry. Although the ABCD Study aimed to achieve its representativeness by recruiting from an array of school systems located around each of the 21 research sites, chosen for their diversity in geography, demographics, and socioeconomic status, it is not fully representative of the US population (Compton et al., 2019). Third, the duration of the follow-up period utilized in this study is relatively short (1-year and 2-year follow-up), which may limit the interpretability of our findings for understanding cognitive and psychiatric development during later childhood. Future research could potentially benefit from employing longer follow-up periods, as more follow-up observations are being collected in the ABCD Study. Fourth, while we used a wide range of statistical methods to adjust for confounding bias from observed and unobserved variables (e.g., genetic ancestry), we did not account for other types of potential bias such as sample selection bias. Fifth, despite a number of causal inference methods used in this study, the ABCD Study is a non-randomized dataset. Given the observational nature of the ABCD Study, interpreting our results as actual causality requires more caution. Finally, we did not include all important environmental variables, such as air pollution (Marshall et al., 2020) and social capital (Krabbendam, 2005), which are not collected in the ABCD study.

In conclusion, our study provides potential pathways of genetic factors of cognitive phenotypes and environmental factors of family, school, and neighborhood to cognitive and mental wellness in children. Our findings underscore the importance of a comprehensive approach that considers both biological and socioeconomic features in promoting young children’s cognitive ability and mental health. Given the importance of child development, it requires joint efforts from multiple disciplines.

Added in the revised version: Supporting information

Removed in the revised version: Funding and Disclosure

Removed in the revised version: This work was supported by the New Faculty Startup Fund from Seoul National University, a Research grant from Basic Science Research Program through the National Research Foundation of Korea (NRF) funded by the Ministry of Education, South Korea (2021R1I1A1A01054995), and Institute of Information & communications Technology Planning & Evaluation (IITP) grant funded by the Korea government (MSIT) [NO.2021-0- 01343, Artificial Intelligence Graduate School Program (Seoul National University)]. The authors have nothing to disclose.

Removed in the revised version: Author Contributions

Removed in the revised version: JP, EL, and GC conducted the main analyses. HH, YYJ, and JC supervised the study. JP, EL preprocessed data. JP, EL, GC, HH performed statistical analyses. EL, BK, GK, and YYJ constructed PGS. JP, EL, GC, YYJ, and JC drafted and critically reviewed the manuscript. All authors read and approved the final draft.

Removed in the revised version: Data Availability

Removed in the revised version: All codes used in this study can be found at https://github.com/Transconnectome/Intell_PLE_Pathway. Due to ABCD Study’s policy in data sharing, we provide a synthetic dataset instead of real observations. This synthetic dataset was generated with conditional GAN (Xu et al., 2019) to imitate the data structure of our final study samples. After automatic hyperparameter optimization with Optuna (Akiba et al., 2019), the synthetic dataset showed overall quality score of 84.15%. Note that analyses results from the synthetic data may not be 100% identical to the results presented in this paper, due to the differences between synthetic vs original dataset.